```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

## Midterm Redemption


### Question 0

The exam redemption consists of 8 questions worth a total of 25 points. Please complete the exam redemption. Here are a few things that differ from the midterm:

Instead of assigning variables directly using the correct code, we convert them into functions. Each function takes no arguments but should return the code you would normally assign to the variable. Here’s an example:
In the exam, you wrote:
```python
answer = 1 + 1
```

But in the redemption, you will write:

```python
def answer_function():
    return 1 + 1
```

Please refer to the lab.py file to see the names of the functions corresponding to each question.

Feel free to post any logistical questions, but course staff will not answer questions related to the answers.

Start by running the cell below to import the packages needed for the exam.

```python
%load_ext autoreload
%autoreload 2
```

```python
from lab import *
```

```python
import pandas as pd
import numpy as np
import sys
import os
from dsc80_utils import *

```

## Parks!
Sam is working with a dataset on visits to US National Parks. There are two CSV files that he's working with: `data/parks.csv` and `data/visits.csv`. The columns of each dataset are described below.

`data/parks.csv`:

- `park_id` (`int`): A unique identifier for a park.
- `parkname` (`str`): The name of each park. `NP` is short for National Park. `PRES` is short for Preserve.
- `region` (`str`): The region of the US that the park is in. 
- `state` (`str`): The two-letter abbreviation for the state that the park is in.

`data/visits.csv`:

- `name` (`str`): The full name of the park visitor. You can assume that no two people share the same first and list name.
- `entry` (`datetime64`): The timestamp that the visitor entered the park.
- `duration` (`int`): The number of hours that each person stayed at the park.
- `park_id` (`int`): The park ID that the person stayed at.

Run the cell below to load the two CSV files into two Pandas DataFrames.

```python
parks = pd.read_csv('data/parks.csv')
visits = pd.read_csv('data/visits.csv', parse_dates=['entry'])
```

```python
parks.head(2)
```

```python
visits.head(2)
```

Now, run the cell below to define two more dataframes.

```python
df1 = parks.merge(visits, on='park_id', how='inner')
df2 = parks.merge(visits, on='park_id', how='left')
```

Write Python code below to compute each of the following desired results using the dataframes `parks`, `visits`, `df1`, and `df2`. Unless otherwise stated, all number and boolean results should be numpy types (e.g. `np.float64` instead of `float`).

### Question 1 
**(3 points)**
Find the number of times that Amina Garcia visited a national park.
Store your result in the variable `amina` as a Python `int`.

```python

```

```python

```

```python
grader.check("q1")
```

### Question 2
**(4 points)**
Find the year with the least number of park visits in this dataset.
Store your result in the variable `least`.

```python

```

```python

```

```python
grader.check("q2")
```

### Question 3 
**(5 points)**
Find the names of all national parks that were **never** visited by someone in this dataset.
Store your result in the variable `never` as a `numpy` array of strings.
(If all national parks were visited, `never` should be an empty numpy array.)

```python

```

```python

```

```python
grader.check("q3")
```

## Dog-gone data

Sam was analyzing the parks data when his dog Bentley ate his data (again!).
Sam's table is stored in a dataframe called `df` that is now missing some values in the `duration` column.
Your task is to impute the missing duration values as effectively as possible, such that the mean and the variance of the `duration` column after imputation is as close as possible to the original data.
First, start by running the cell below to load in Sam's data.

```python
df = pd.read_csv('data/visits_missing_a.csv')
df
```

## Question 4
**(1 point)**
What proportion of values are missing in the `duration` column?
Store your result in the variable `prop_missing` as a number between 0 and 1.

```python

```

```python

```

```python
grader.check("q4")
```

### Question 5 
**(2 points)**
We've imported the following functions from `dsc80_utils.py`. 

- `permutation_test(data, col, group_col, test_statistic, N)`
- `diff_in_means(data, col, group_col)`
- `tvd(data, col, group_col)`
- `ks(data, col, group_col)`

As usual, you can view the complete docstrings for these functions by adding a `?` after the function name, for example, by running:

```python
permutation_test?
```

Now, run the following cell, which runs a permutation test for the following hypotheses:

- Null: The distribution of `entry_hour` is the same for park entries in 2020 and the rest of the years.
- Alternative: The distribution of `entry_hour` is different for park entries in 2020.
Using the results from `permutation_test`, compute the p-value and store it in `pval`. Then, set `reject` to `True` if you reject the null hypothesis or `False` if you fail to reject the null at the 0.01 significance level.

```python

```

```python

```

```python
grader.check("q5")
```

### Question 6
**(3 points)** Run a permutation test to determine whether the missingness of `duration` is missing at random (MAR) dependent on `entry_hour`. 
Store your p-value in the variable `hour_pval`, and your conclusion in `hour_mar` at the 0.01 significance level (`True` if `duration` is MAR dependent on `entry_hour` and `False` if not). Be sure to choose the appropriate test statistic for this hypothesis test.

```python

```

```python

```

```python
grader.check("q6")
```

### Question 7
**(4 points)** Now, run a permutation test to determine whether the missingness of `duration` is missing at random (MAR) dependent on **any other column** in `df`: `region`, `entry_year`, and `entry_month`. Use a 0.01 significance level.

For each permutation test, choose the appropriate test statistic, and store your p-values in the dictionary `pvals`. The dictionary should have one key for each column (`region`, `entry_year`, and `entry_month`). Each value in the dictionary should be the p-value for the hypothesis test to determine whether `duration` is MAR on that column.

Then, mark each column that `duration` is MAR on in the dictionary `mar_cols` (`True` if `duration` is MAR on that column and `False` if not).

```python

```

```python

```

```python
## don't delete this cell, but do run it -- it is needed for the autograder tests
mar_cols = mar_cols_function()
pval = pvals_functions()
```

```python
grader.check("q7")
```

## Question 8
**(3 points)** Impute the missing `duration` values using probabilistic imputation, conditional on the columns that `duration` is MAR on.
You should make use of the `prob_impute` function defined below (which is copied directly from lecture).

Store your result in the variable `df_imputed`, which is a copy of `df` with one extra column called `imputed` containing the imputed duration values.

```python
def prob_impute(s):
    s = s.copy()
    num_null = s.isna().sum()
    fill_values = np.random.choice(s.dropna(), num_null)
    s[s.isna()] = fill_values
    return s
def df_imputed_function():
    ...
```

```python

```

```python

```

```python
## don't delete this cell, but do run it -- it is needed for the autograder tests
df_imputed = df_imputed_function()
```

```python
grader.check("q8")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: None)
```
>>> amina = amina_function()
>>> type(amina) == int
True
```


Tests for q2 (points: None)
```
>>> bool(least_function() > 2000)
True
```
```
>>> #visits['entry'].dt.year.value_counts()
```


Tests for q3 (points: None)
```
>>> isinstance(never_function(), np.ndarray)
True
```
```
>>> all(isinstance(i, str) for i in never_function())
True
```


Tests for q4 (points: None)
```
>>> isinstance(prop_missing_function(), np.floating)
True
```
```
>>> bool(0 < prop_missing_function() < 1)
True
```


Tests for q5 (points: None)
```
>>> isinstance(p_val_function(), np.floating)
True
```
```
>>> isinstance(reject_function(), bool)
True
```


Tests for q6 (points: None)
```
>>> isinstance(hour_pval_function(), np.floating)
True
```
```
>>> isinstance(hour_mar_function(), bool)
True
```


Tests for q7 (points: None)
```
>>> isinstance(pval, dict)
True
```
```
>>> set(pval.keys()) == {'region', 'entry_year', 'entry_month'}
True
```
```
>>> all(isinstance(i, np.floating) for i in pval.values())
True
```
```
>>> all(0 <= i <= 1 for i in pval.values())
True
```
```
>>> isinstance(mar_cols, dict)
True
```
```
>>> set(mar_cols.keys()) == {'region', 'entry_year', 'entry_month'}
True
```
```
>>> all(isinstance(i, (bool, np.bool_)) for i in mar_cols.values())
True
```


Tests for q8 (points: None)
```
>>> isinstance(df_imputed, pd.DataFrame)
True
```
```
>>> set(df_imputed.columns) == {'region', 'duration', 'entry_year', 'entry_month', 'entry_hour', 'imputed'}
True
```
```
>>> df_imputed.shape == (5000, 6)
True
```
